
<span style="font-size:x-large;">Lab 3</span>

# Exploring the Census Age Data

In *The Census and the "Typical" Australian?"*, we discussed the age chosen for Clare (38), the "Typical" Australian, and Paul (37), the "Typical" Western Australian.

We saw that the ABS chose to use the *mode* for most attributes, while for age they chose to use the *median*. We asked why they may have made that decision. (Noting as well that in their Quickstats they also used some *means* (averages), somewhat confusingly, for the same attributes.)

In this lab we'll explore the data, and see if you think that choice is justified.

*It is recommended that the case study videos are viewed in conjunction with this lab - either beforehand to get some direction and/or afterwards to review your approach, depending on your learning strategy.*

## Data Acquisition

Download the 2016 Census *General Community Profile* datapack from the ABS.

Open the Metadata spreadsheet (in your preferred spreadsheet software) and find the table that reports age data from the Census.

Go to the Cell Descriptors for that data. You should find age data starting from the age of zero.

_Q: What happens when you get to 80?_

You should see that people over 80 do not get the same representation as those under 80. Can you think of any good reason to justify this blatant ageism?

> _Before you upload the data to CoCalc, check the file size. If you choose to upload the whole zip file, it should be less than 5MB (if not, check that you have only downloaded the data for Australia as a whole, not all the separate regions as well.) It is recommended that you remove an additional 3.5MB by unzipping the package and deleting the file `2016Census_geog_desc_1st_and_2nd_release.xlsx` which is not needed for the labs._
>
> _Alternatively, you can just upload the components you need. The tables (csv files) you need for the lab are each less than 5KB._


* Upload the data to the directory for this Lab.

* Set up suitable constants to point to the age data. Read in and print out the data to ensure you have accessed it correctly.

```python
DATA_A= "2016Census_G04A_AUS.csv"
DATA_B= "2016Census_G04B_AUS.csv"
```

## Data Cleaning

We are going to examine the age profile of the *population as a whole*.

Have another look at the Cell Descriptors in the spreadsheet. Which rows of data do you need for this task?

Let's start by pulling the data into lists.

* Write some code to read in the data and store the data you need for this task in lists. You can strip out some of the rows you don't need at this stage, however don't strip out the 5 year totals yet, as we'll need those for the next step. Make any appropriate type conversions, and tidy up category names if you wish.

Tip: A similar example is covered in the videos.

For comparison, I have two lists with 101 items in each after this step.


```python
DATA_A= "2016Census_G04A_AUS.csv"


with open(DATA_A) as file:
    Age_Category_A = file.readline()[1:].split(',')[1:-2]
    Age_Numbers_A = file.readline()[1:].split(',')[1:-2]
    print(Age_Category_A)
    print(Age_Numbers_A)

   
    
```

Output:
```
['Age_yr_0_M', 'Age_yr_0_F', 'Age_yr_0_P', 'Age_yr_1_M', 'Age_yr_1_F', 'Age_yr_1_P', 'Age_yr_2_M', 'Age_yr_2_F', 'Age_yr_2_P', 'Age_yr_3_M', 'Age_yr_3_F', 'Age_yr_3_P', 'Age_yr_4_M', 'Age_yr_4_F', 'Age_yr_4_P', 'Age_yr_0_4_M', 'Age_yr_0_4_F', 'Age_yr_0_4_P', 'Age_yr_5_M', 'Age_yr_5_F', 'Age_yr_5_P', 'Age_yr_6_M', 'Age_yr_6_F', 'Age_yr_6_P', 'Age_yr_7_M', 'Age_yr_7_F', 'Age_yr_7_P', 'Age_yr_8_M', 'Age_yr_8_F', 'Age_yr_8_P', 'Age_yr_9_M', 'Age_yr_9_F', 'Age_yr_9_P', 'Age_yr_5_9_M', 'Age_yr_5_9_F', 'Age_yr_5_9_P', 'Age_yr_10_M', 'Age_yr_10_F', 'Age_yr_10_P', 'Age_yr_11_M', 'Age_yr_11_F', 'Age_yr_11_P', 'Age_yr_12_M', 'Age_yr_12_F', 'Age_yr_12_P', 'Age_yr_13_M', 'Age_yr_13_F', 'Age_yr_13_P', 'Age_yr_14_M', 'Age_yr_14_F', 'Age_yr_14_P', 'Age_yr_10_14_M', 'Age_yr_10_14_F', 'Age_yr_10_14_P', 'Age_yr_15_M', 'Age_yr_15_F', 'Age_yr_15_P', 'Age_yr_16_M', 'Age_yr_16_F', 'Age_yr_16_P', 'Age_yr_17_M', 'Age_yr_17_F', 'Age_yr_17_P', 'Age_yr_18_M', 'Age_yr_18_F', 'Age_yr_18_P', 'Age_yr_19_M', 'Age_yr_1
```

```python
print(len(Age_Category_A), len(Age_Numbers_A))
```

Output:
```
198 198
```

```python

DATA_B= "2016Census_G04B_AUS.csv"

with open(DATA_B) as file:
    Age_Category_B = file.readline()[1:].split(',')[1:-3]
    Age_Numbers_B = file.readline()[1:].split(',')[1:-3]
    print(Age_Category_B)
    print(Age_Numbers_B)
```

Output:
```
['Age_yr_55_P', 'Age_yr_56_M', 'Age_yr_56_F', 'Age_yr_56_P', 'Age_yr_57_M', 'Age_yr_57_F', 'Age_yr_57_P', 'Age_yr_58_M', 'Age_yr_58_F', 'Age_yr_58_P', 'Age_yr_59_M', 'Age_yr_59_F', 'Age_yr_59_P', 'Age_yr_55_59_M', 'Age_yr_55_59_F', 'Age_yr_55_59_P', 'Age_yr_60_M', 'Age_yr_60_F', 'Age_yr_60_P', 'Age_yr_61_M', 'Age_yr_61_F', 'Age_yr_61_P', 'Age_yr_62_M', 'Age_yr_62_F', 'Age_yr_62_P', 'Age_yr_63_M', 'Age_yr_63_F', 'Age_yr_63_P', 'Age_yr_64_M', 'Age_yr_64_F', 'Age_yr_64_P', 'Age_yr_60_64_M', 'Age_yr_60_64_F', 'Age_yr_60_64_P', 'Age_yr_65_M', 'Age_yr_65_F', 'Age_yr_65_P', 'Age_yr_66_M', 'Age_yr_66_F', 'Age_yr_66_P', 'Age_yr_67_M', 'Age_yr_67_F', 'Age_yr_67_P', 'Age_yr_68_M', 'Age_yr_68_F', 'Age_yr_68_P', 'Age_yr_69_M', 'Age_yr_69_F', 'Age_yr_69_P', 'Age_yr_65_69_M', 'Age_yr_65_69_F', 'Age_yr_65_69_P', 'Age_yr_70_M', 'Age_yr_70_F', 'Age_yr_70_P', 'Age_yr_71_M', 'Age_yr_71_F', 'Age_yr_71_P', 'Age_yr_72_M', 'Age_yr_72_F', 'Age_yr_72_P', 'Age_yr_73_M', 'Age_yr_73_F', 'Age_yr_73_P', 'Age_yr_74_M
```

```python
print(len(Age_Category_B),len(Age_Numbers_B))
```

Output:
```
103 103
```

```python
Age_A = Age_Category_A[slice(2,len(Age_Category_A),3)]
Age_Pop_A = Age_Numbers_A[slice(2,len(Age_Numbers_A),3)]

print(len(Age_A), len(Age_Pop_A))
```

Output:
```
File "<ipython-input-19-8a9736ed2494>", line 1
    Age_A = Age_Category_A[slice(1:9,len(Age_Category_A),0)]
                                  ^
SyntaxError: invalid syntax
```

```python

```

```python
Age_B = Age_Category_B[slice(0,len(Age_Category_B),3)]
Age_Pop_B = Age_Numbers_B[slice(0,len(Age_Numbers_B),3)]

print(len(Age_B), len(Age_Pop_B))
```

Output:
```
35 35
```

```python
Age_A
```

Output:
```
['Age_yr_0_P',
 'Age_yr_1_P',
 'Age_yr_2_P',
 'Age_yr_3_P',
 'Age_yr_4_P',
 'Age_yr_0_4_P',
 'Age_yr_5_P',
 'Age_yr_6_P',
 'Age_yr_7_P',
 'Age_yr_8_P',
 'Age_yr_9_P',
 'Age_yr_5_9_P',
 'Age_yr_10_P',
 'Age_yr_11_P',
 'Age_yr_12_P',
 'Age_yr_13_P',
 'Age_yr_14_P',
 'Age_yr_10_14_P',
 'Age_yr_15_P',
 'Age_yr_16_P',
 'Age_yr_17_P',
 'Age_yr_18_P',
 'Age_yr_19_P',
 'Age_yr_15_19_P',
 'Age_yr_20_P',
 'Age_yr_21_P',
 'Age_yr_22_P',
 'Age_yr_23_P',
 'Age_yr_24_P',
 'Age_yr_20_24_P',
 'Age_yr_25_P',
 'Age_yr_26_P',
 'Age_yr_27_P',
 'Age_yr_28_P',
 'Age_yr_29_P',
 'Age_yr_25_29_P',
 'Age_yr_30_P',
 'Age_yr_31_P',
 'Age_yr_32_P',
 'Age_yr_33_P',
 'Age_yr_34_P',
 'Age_yr_30_34_P',
 'Age_yr_35_P',
 'Age_yr_36_P',
 'Age_yr_37_P',
 'Age_yr_38_P',
 'Age_yr_39_P',
 'Age_yr_35_39_P',
 'Age_yr_40_P',
 'Age_yr_41_P',
 'Age_yr_42_P',
 'Age_yr_43_P',
 'Age_yr_44_P',
 'Age_yr_40_44_P',
 'Age_yr_45_P',
 'Age_yr_46_P',
 'Age_yr_47_P',
 'Age_yr_48_P',
 'Age_yr_49_P',
 'Age_yr_45_49_P',
 'Age_yr_50_P',
 'Age_yr
```

```python
Age = Age_A + Age_B
Population = Age_Pop_A + Age_Pop_B

```

```python
DATA_A= "2016Census_G04A_AUS.csv"
DATA_B= "2016Census_G04B_AUS.csv"

with open(DATA_A) as file:
    Age_Category_A = file.readline()[1:].split(',')[1:-2]
    Age_Numbers_A = file.readline()[1:].split(',')[1:-2]

with open(DATA_B) as file:
    Age_Category_B = file.readline()[1:].split(',')[1:-3]
    Age_Numbers_B = file.readline()[1:].split(',')[1:-3]

Age_A = Age_Category_A[slice(2,len(Age_Category_A),3)]
Age_Pop_A = Age_Numbers_A[slice(2,len(Age_Numbers_A),3)]

Age_B = Age_Category_B[slice(0,len(Age_Category_B),3)]
Age_Pop_B = Age_Numbers_B[slice(0,len(Age_Numbers_B),3)]

Age = Age_A + Age_B
Population = Age_Pop_A + Age_Pop_B

print(len(Age),len(Population))


```

Output:
```
101 101
```

```python
for i in range(len(Population)):
    Population[i] = int(Population[i])
    
```

```python
type(Population[0])
```

Output:
```
int
```

```python
print(len(Age),len(Population))
```

Output:
```
101 101
```

```python
for i in range(len(Age)):
    if '_P' in Age[i]:
        Age[i] = Age[i][:-2]
        

```

```python
type(Age),type(Population)
```

Output:
```
(list, list)
```

```python
type(Age[0]),type(Population[0])
```

Output:
```
(str, int)
```

```python
DATA_A= "2016Census_G04A_AUS.csv"
DATA_B= "2016Census_G04B_AUS.csv"

with open(DATA_A) as file:
    Age_Category_A = file.readline()[1:].split(',')[1:-2]
    Age_Numbers_A = file.readline()[1:].split(',')[1:-2]

with open(DATA_B) as file:
    Age_Category_B = file.readline()[1:].split(',')[1:-3]
    Age_Numbers_B = file.readline()[1:].split(',')[1:-3]

Age_A = Age_Category_A[slice(2,len(Age_Category_A),3)]
Age_Pop_A = Age_Numbers_A[slice(2,len(Age_Numbers_A),3)]

Age_B = Age_Category_B[slice(0,len(Age_Category_B),3)]
Age_Pop_B = Age_Numbers_B[slice(0,len(Age_Numbers_B),3)]

Age = Age_A + Age_B
Population = Age_Pop_A + Age_Pop_B

for i in range(len(Population)):
    Population[i] = int(Population[i])

for i in range(len(Age)):
    if '_P' in Age[i]:
        Age[i] = Age[i][:-2]
        
```

```python
Age
```

Output:
```
['Age_yr_0',
 'Age_yr_1',
 'Age_yr_2',
 'Age_yr_3',
 'Age_yr_4',
 'Age_yr_0_4',
 'Age_yr_5',
 'Age_yr_6',
 'Age_yr_7',
 'Age_yr_8',
 'Age_yr_9',
 'Age_yr_5_9',
 'Age_yr_10',
 'Age_yr_11',
 'Age_yr_12',
 'Age_yr_13',
 'Age_yr_14',
 'Age_yr_10_14',
 'Age_yr_15',
 'Age_yr_16',
 'Age_yr_17',
 'Age_yr_18',
 'Age_yr_19',
 'Age_yr_15_19',
 'Age_yr_20',
 'Age_yr_21',
 'Age_yr_22',
 'Age_yr_23',
 'Age_yr_24',
 'Age_yr_20_24',
 'Age_yr_25',
 'Age_yr_26',
 'Age_yr_27',
 'Age_yr_28',
 'Age_yr_29',
 'Age_yr_25_29',
 'Age_yr_30',
 'Age_yr_31',
 'Age_yr_32',
 'Age_yr_33',
 'Age_yr_34',
 'Age_yr_30_34',
 'Age_yr_35',
 'Age_yr_36',
 'Age_yr_37',
 'Age_yr_38',
 'Age_yr_39',
 'Age_yr_35_39',
 'Age_yr_40',
 'Age_yr_41',
 'Age_yr_42',
 'Age_yr_43',
 'Age_yr_44',
 'Age_yr_40_44',
 'Age_yr_45',
 'Age_yr_46',
 'Age_yr_47',
 'Age_yr_48',
 'Age_yr_49',
 'Age_yr_45_49',
 'Age_yr_50',
 'Age_yr_51',
 'Age_yr_52',
 'Age_yr_53',
 'Age_yr_54',
 'Age_yr_50_54',
 'Age_yr_55',
 'Age_yr_56',
 'Age_yr_57',
 'Age_yr_58',
 
```

### Data Augmentation


As the ABS have inexplicably left off people over 80, we will have to reconstruct the data as best we can. This is akin to the *Statistical Imputation* discussed in the videos (where we discussed Hot Decking, Mean Substitution, Probabilistic and Regression methods).

#### Checked Answer [1 prac mark]

As a first approximation we will simply use an averaging approach. We'll make an *assumption* that the number of people in a five year category over 80 is the approximately the same for each year.

* Write a function `spread(age_cat, num)` that:
  * takes an argument `age_cat` that represents a date range in the format `Age_yr_`*n1_n2* (eg. `Age_yr_80_84`) and an integer `num`
  * returns a pair (2-tuple) of lists:
    * the first list should contain strings in the form `Age_yr_`*n* where *n* ranges from the first to the last year in the range (eg. `Age_yr_80`, `Age_yr_81`,...)
    * the second list contain a list of integers adding up to `num` such that:
      * if `num` is divisible by the number of years, then each integer should be the same (ie. the people are evenly spread)
      * if there is a remainder from dividing `num` by the number of years, the remainder should be added one at a time to the year groups starting at the lowest age (80 in this example) until there are no more remaining

For example, `spread('Age_yr_80_84', 7)` should return:
```
(['Age_yr_80', 'Age_yr_81', 'Age_yr_82', 'Age_yr_83', 'Age_yr_84'], [2, 2, 1, 1, 1])
```

Tip: You might find `divmod` useful. Also, recall that `+` and `*` can be used to create lists.


<!-- grade_id: spread, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal
assert_equal(spread('Age_yr_80_84', 7), (['Age_yr_80', 'Age_yr_81', 'Age_yr_82', 'Age_yr_83', 'Age_yr_84'], [2, 2, 1, 1, 1]))
assert_equal(spread('Age_yr_80_89', 7), (['Age_yr_80', 'Age_yr_81', 'Age_yr_82', 'Age_yr_83', 'Age_yr_84', 'Age_yr_85', 'Age_yr_86', 'Age_yr_87', 'Age_yr_88', 'Age_yr_89'], [1, 1, 1, 1, 1, 1, 1, 0, 0, 0]))
assert_equal(spread('Age_yr_80_84', 0), (['Age_yr_80', 'Age_yr_81', 'Age_yr_82', 'Age_yr_83', 'Age_yr_84'], [0, 0, 0, 0, 0]))
assert_equal(spread('Age_yr_80_84', 460549), (['Age_yr_80', 'Age_yr_81', 'Age_yr_82', 'Age_yr_83', 'Age_yr_84'], [92110, 92110, 92110, 92110, 92109]))
print("So far, so good. Additional test cases will be applied.")

```

Output:
```
So far, so good. Additional test cases will be applied.
```

* Using your `spread` function, complete (reconstruct) the lists of age categories, following the pattern of lower age years, as follows:
  * for the age ranges from 80-84 to 95-99, insert the categories and ages for each year calculated by spreading the total number over the previous 5 years
  * for the age range 100 years and over, assume the ages are spread over the range 100-104 (again a poor assumption) and insert as above

* Combine your list of categories and your list of numbers into a list of pairs containing *only entries for each single year*.
* In the process also strip out "`yr_`" from the category names.

Tip: Use `zip` to combine the lists.

Your list should now start like this:
```
[('Age_0', 276227),
 ('Age_1', 293503),
 ('Age_2', 295142),
 ('Age_3', 299725),
 ('Age_4', 300184),
 ('Age_5', 298271),
 ('Age_6', 302901),
 ('Age_7', 299413),
 ...
```

* Check your list has the number of entries that you would expect.


#### Checked Answer [1 prac mark]

* Drawing from your work so far, define a function `cleaned_data()` that:
    * reads in the age data from the ABS datapack
    * augments the data with the missing age data as described above
    * keeps only the necessary data (number of persons for each year)
    * shortens the category names to the form `Age_`*n*
    * returns the data in a list of pairs in the form (*category_name*, *number*)

Check that your output is as you would expect, before applying the sample tests.

Note that your function can call functions you have defined previously. As usual, however, check that it works with a fresh kernel.

<!-- grade_id: cleaned_data, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal
assert_equal(cleaned_data()[0], ('Age_0', 276227))
assert_equal(cleaned_data()[-1], ('Age_105', 594))
assert_equal(len(cleaned_data()), 106)
print("So far, so good. Additional test cases will be applied.")

```

Output:
```
So far, so good. Additional test cases will be applied.
```

### Checked Answer [1 prac mark]

* Without using any libraries, write a function `central_measures(clean_data)` that takes a list of (age_year, number) pairs, and returns a triple (mean, median, mode), with mean as the nearest integer, containing the measures of central tendency for age for the list.

Test your function thoroughly. Some example tests are included below. Do you agree with the answers given?



<!-- grade_id: central_measures, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal
assert_equal(central_measures(cleaned_data()[:1]), (0,0,0))
assert_equal(central_measures(cleaned_data()[:2]), (1,1,1))
assert_equal(central_measures(cleaned_data()[:3]), (1,1,2))
assert_equal(central_measures(cleaned_data()[:4]), (2,2,3))

```

_Q: What are the measures of central tendency for your data as a whole? How do they compare with the "typical" Australians?_

Bear in mind that these numbers will depend on the assumptions we made when filling in the missing data!

## Visualisation

* Use `pyplot` to produce a basic vertical bar chart for your data.

How would you improve your chart to better understand the data from the plot?

* Improve your plot to make it 'production-ready':
  * choose an appropriate width for the plot that allows individual years to be better distinguished
  * place the ticks at 5 year intervals along the x-axis
  * label the ticks, with the labels written vertically
  * label your axes and give the chart a title
  * save your chart to a png file

The bottom of your plot should look like this:

<img src="bottom-of-plot.png" width=1000>

Open the png file in CoCalc and check that it looks as you would expect.

_Q: What can you learn from the plot? What do you make of the ABS' assertion that the 'typical Australians' are in the 36-38 year range? What age would you choose looking at the distribution?_

_Q: Peter Costello introduced the 'baby bonus' around 2002, and made his plea to parents in 2004. The census was taken in 2016. Do you see any (circumstantial) evidence that parents may have heeded his plea? (Bearing in mind, of course, that correlation does not imply causation.))_

## Challenge

We filled in the missing data using the average for the age range (in a similar spirit to Mean Substitution). This hasn't turned out too badly, but we could do better.

* While staying true to the total number in each age range, improve your method for filling in the missing data to make it as realistic as you can. Remember that you would like to try to use a *general* approach that would also work on other sets of data, rather than one that is *specific* to this set of data.

(Fitting the data to a trend is a bit like using a Regression approach to Statistical Imputation, although you are not being asked to use regression here. We will come to this later.)

Which of the measures of central tendency could be affected by this? Check whether the measures have changed.

## More Babies

*The Census and the 'Typical' Australian?, Part 6* includes an exercise to change the line plot so that it uses the percentage on the y-axis rather than the raw numbers. *Part 7* asks if you see evidence for the Treasurer's baby boom.

Go ahead and complete the exercise if you haven't done so already.

&copy; Cara MacNish