<span style="font-size:x-large;">Lab 2 (*Fake News!* Case Study)</span>

# Covid Tests per Capita: Is the US Leading the World?

![](trump-testing.jpg)

President Trump has repeatedly said that the US tests more than any other country in the world by far, and sometimes more than all the other countries put together. See for example the official Whitehouse transcript: [Remarks by President Trump on Supporting our Nation’s Small Businesses Through the Paycheck Protection Program](https://www.whitehouse.gov/briefings-statements/remarks-president-trump-supporting-nations-small-businesses-paycheck-protection-program/), 28th April, 2020.

In this practical work, as part of our *Fake News!* case study, we'll investigate the veracity of one aspect of these claims, the *per capita* test rates.

## Data Acquisition

#### *Our World in Data*

In the videos for this case study we obtained data, originally from the European CDC, through *Our World in Data*: https://ourworldindata.org/.

* Read the *[About](https://ourworldindata.org/about)* page to find out what this initiative is, what they hope to achieve, and who it is backed by. It is always important to know who is behind a data source in order to make an assessment about what degree of credibility to give the source.

Follow the link to `Health | Coronavirus Pandemic` and then to `Tests`. 

Note that the Data Scientists behind this site have provided a number of different ways of looking at and interpreting test rates. There is also a lot of background provided. 

* Scroll down to the section "Our checklist for COVID-19 testing data" and read through the ten items on the checklist.

This is terrific example of Data Science done well!




### Per-capita testing

Find the section entitled "How many tests are performed each day".

Have a look at how the Map is presented, with the time slider that allows you to see snapshots over time.

Then look at the Chart view and scroll the cursor over the chart.

Again, these are great examples of interactive data presentation.

Note: Consistent with many of the media sites, we will refer to these data in general terms as the "per capita" data to distinguish them from the totals data. More precisely, however, they are tests per 1000 people. That is, they are exactly 1000 times the per capita rate. The only reason for multiplying by 1000 is that its easier to read, say, 0.08 than 0.00008.

### Downloading and uploading the data

* Go to `DOWNLOAD` and download the csv file `daily-tests-per-thousand-people-smoothed-7-day.csv`.

* Then open the folder icon in CoCalc, make sure you are in the same directory as this lab sheet, and drop (upload) the csv file into the directory.

Click on the file to open it in CoCalc and have a look at the file format.

You should see, as anticipated from the filename extension, that this is a *comma separated values (csv)* file: in each row, the *fields* are separated by commas. There is also a header line, indicating what the data in each field represents.


### Reading in the data

As usual, start by setting up a constant with the path (in this case, its just the name) of the data file, so you don't need to keep typing it. For this lab we'll use the version of this file from the 29th July, distributed with the lab, so that we are all using the same file. (Feel free to run your code on your own version too - just be aware that the outputs and test results may be different to those in this sheet.)

You can access this file with:

`DATA = "daily-tests-per-thousand-people-smoothed-7-day-20200729.csv"`

Note that in this lab we won't put empty cells for you to complete your code in. You can create cells as you need them using the '+' button.

* Read and print out the first 5 lines of data. The output should start like this:

```
Entity,Code,Date,Daily tests per thousand people (7-day smoothed) (tests per thousand)

Argentina,ARG,"Feb 18, 2020",0

Argentina,ARG,"Feb 19, 2020",0
```


```python
DATA = "daily-tests-per-thousand-people-smoothed-7-day-20200729.csv"
```

```python
with open(DATA) as file:
    for line in range(5):
        print(file.readline())
```

Output:
```
Entity,Code,Date,Daily tests per thousand people (7-day smoothed) (tests per thousand)

Argentina,ARG,"Feb 18, 2020",0

Argentina,ARG,"Feb 19, 2020",0

Argentina,ARG,"Feb 20, 2020",0

Argentina,ARG,"Feb 21, 2020",0
```

```python

```

## Data Conversion and Cleaning

#### From strings to lists

* Read the first 5 lines again. This time use the `split` method to turn each line into a list before printing it out.

Your output should start like this:
```
['Entity', 'Code', 'Date', 'Daily tests per thousand people (7-day smoothed) (tests per thousand)\n']
['Argentina', 'ARG', '"Feb 18', ' 2020"', '0\n']
['Argentina', 'ARG', '"Feb 19', ' 2020"', '0\n']
```

```python
with open(DATA) as file:
    for line in range(5):
        print(file.readline().split())
```

Output:
```
['Entity,Code,Date,Daily', 'tests', 'per', 'thousand', 'people', '(7-day', 'smoothed)', '(tests', 'per', 'thousand)']
['Argentina,ARG,"Feb', '18,', '2020",0']
['Argentina,ARG,"Feb', '19,', '2020",0']
['Argentina,ARG,"Feb', '20,', '2020",0']
['Argentina,ARG,"Feb', '21,', '2020",0']
```

Notice that we still have the newline character in the last item.

* Use the `strip` function to remove whitespace before splitting the lines.

Your output should now start like this:
```
['Entity', 'Code', 'Date', 'Daily tests per thousand people (7-day smoothed) (tests per thousand)']
['Argentina', 'ARG', '"Feb 18', ' 2020"', '0']
['Argentina', 'ARG', '"Feb 19', ' 2020"', '0']
```

```python
with open(DATA) as file:
    for line in range(5):
        print(file.readline().strip().split())
```

Output:
```
['Entity,Code,Date,Daily', 'tests', 'per', 'thousand', 'people', '(7-day', 'smoothed)', '(tests', 'per', 'thousand)']
['Argentina,ARG,"Feb', '18,', '2020",0']
['Argentina,ARG,"Feb', '19,', '2020",0']
['Argentina,ARG,"Feb', '20,', '2020",0']
['Argentina,ARG,"Feb', '21,', '2020",0']
```

You may have noticed another problem, caused by the fact that the dates include commas. In fact, you may conclude that with this date format, the choice of a comma as the *delimiter* (separator) was not a particularly good one, and an alternative such as tab separated (tsv) would have been better choice for these data.

Nevertheless, it is not ambiguous, because commas that are not intended as delimiters only appear within double quotes. The 'user' (our code in this case) is expected to take the quotes into account when splitting the lines. 

There are many ways to deal with dates, but for now, we can just remove the comma and the quotes, since neither provide us with any information. The fields within the quotes (month, day and year) can be distinguished by their order (the comma is just for human consumption) and the quotes are redundant since it is a text file and all the fields will be read as strings.

Change your code so that it has a *preprocessing* step before it splits the lines. For each line after the header line your preprocessing step should:
  * find the positions of the two double quotes
  * replace the comma between the quotes with a space
  * replace the old date with the new date without a comma
  * throw a `ValueError` if the line doesn't have double quotes
  * Are there any empty row, if yes, delete empty row

We'll break it down into steps.


```python

```

```python

```

```python

        
with open(DATA) as file:
    for line in file:
        lines = line.strip()
        
        
        while ',,' in lines:
            lines = lines.replace(',,',',')

        while '(' in lines:
            first = lines.index('(')
            second = lines.index(')')
            third = lines.index('(', first +1)
            fourth = lines.index(')', second + 1)
            field = lines[first:fourth+1]
            lines = lines.replace(field, '')
        
        #print(lines)
#         while '"' in lines:
#             first = lines.find('"')
#             second = lines.find('"', first+1)
# #             third = lines.find('"', second +1)
# #             fourth = lines.find('"', third + 1)
#             field = lines[first:fourth+1]
#             clean_field = field.replace('"',"")[0:-1]
#             lines = lines.replace(field, clean_field)

        while '"' in lines:
            first = lines.find('"')
            second = lines.find('"', first+1)
            field = lines[first:second+1]
            clean_field = field.replace(',',"")[1:-1]
            lines = lines.replace(field, clean_field)
        
        print(lines)

```

Output:
```
Entity,Code,Date,Daily tests per thousand people 
Argentina,ARG,Feb 18 2020,0
Argentina,ARG,Feb 19 2020,0
Argentina,ARG,Feb 20 2020,0
Argentina,ARG,Feb 21 2020,0
Argentina,ARG,Feb 22 2020,0
Argentina,ARG,Feb 23 2020,0
Argentina,ARG,Feb 24 2020,0
Argentina,ARG,Feb 25 2020,0
Argentina,ARG,Feb 26 2020,0
Argentina,ARG,Feb 27 2020,0
Argentina,ARG,Feb 28 2020,0
Argentina,ARG,Feb 29 2020,0
Argentina,ARG,Mar 1 2020,0
Argentina,ARG,Mar 2 2020,0
Argentina,ARG,Mar 3 2020,0
Argentina,ARG,Mar 4 2020,0
Argentina,ARG,Mar 5 2020,0
Argentina,ARG,Mar 6 2020,0
Argentina,ARG,Mar 7 2020,0
Argentina,ARG,Mar 8 2020,0
Argentina,ARG,Mar 9 2020,0
Argentina,ARG,Mar 10 2020,0
Argentina,ARG,Mar 11 2020,0
Argentina,ARG,Mar 12 2020,0
Argentina,ARG,Mar 13 2020,0.001
Argentina,ARG,Mar 14 2020,0.001
Argentina,ARG,Mar 15 2020,0.001
Argentina,ARG,Mar 16 2020,0.001
Argentina,ARG,Mar 17 2020,0.002
Argentina,ARG,Mar 18 2020,0.002
Argentina,ARG,Mar 19 2020,0.002
Argentina,ARG,Mar 20 2020,0.003
Argentina,ARG,Mar 21 2020,0.004
```

```python

```

```python


with open(DATA) as file:
    for i in range(5):

        while ',,' in lines:
            lines = lines.replace(',,',',')

        while '(' in lines:
            first = lines.index('(')
            second = lines.index(')')
            third = lines.index('(', first +1)
            fourth = lines.index(')', second + 1)
            field = lines[first:fourth+1]
            lines = lines.replace(field, '')
            
        while '"' in lines:
            first = lines.find('"')
            second = lines.find('"', first+1)
            third = lines.find('"', second +1)
            fourth = lines.find('"', third + 1)
            field = lines[first:fourth+1]
            clean_field = field.replace('"',"")[0:-1]
            lines = lines.replace(field, clean_field)
        
        print(lines)

        

            
```

Output:
```
Zimbabwe,ZWE,Jul 22 2020,0.048
Zimbabwe,ZWE,Jul 22 2020,0.048
Zimbabwe,ZWE,Jul 22 2020,0.048
Zimbabwe,ZWE,Jul 22 2020,0.048
Zimbabwe,ZWE,Jul 22 2020,0.048
```

```python

```

```python

```

```python

```

```python
with open(DATA) as file:
    for count, line in enumerate(file):
        line = line.strip()
        
        while ',,' in line:
            line = line.replace(",,",',')
     
        while '(' in line:
                first = line.index('(')
                second = line.index(')', first+1)
                third = line.index('(', second + 1)
                fourth = line.index(')', third+1)
                field = line[first-1:fourth+1]
                line = line.replace(field,'')

                
            
        while '"' in line:
            first = line.index('"')
            second = line.index('"', first+1)
            field = line[first:second+1]
            clean_field = field.replace(',',"")[1:-1]
            line = line.replace(field, clean_field)
            
        
        try:
            line.index('""')
        except ValueError:
            raise ValueError('Double Quote Not Found')

```

Output:
```
ValueError                                Traceback (most recent call last)
<ipython-input-7-593e5089513e> in <module>
     27             line.index('""')
     28         except ValueError:
---> 29             raise ValueError('Double Quote Not Found')

ValueError: Double Quote Not Found
```

```python

```

```python

```

```python
def clean(data_row):
    
    content = []
    with open(data_row) as file:
        for line in file:
            lines = line.strip()


            while ',,' in lines:
                lines = lines.replace(',,',',')

            while '(' in lines:
                first = lines.index('(')
                second = lines.index(')')
                third = lines.index('(', first +1)
                fourth = lines.index(')', second + 1)
                field = lines[first:fourth+1]
                lines = lines.replace(field, '')

            while '"' in lines:
                first = lines.find('"')
                second = lines.find('"', first+1)
                third = lines.find('"', second +1)
                fourth = lines.find('"', third + 1)
                field = lines[first:fourth+1]
                clean_field = field.replace('"',"")[0:-1]
                lines = lines.replace(field, clean_field)
            
            content.append(lines)
        return content
    
```

```python


```

```python

def clean(lines):

    while ',,' in lines:
        lines = lines.replace(',,',',')

    while '(' in lines:
        first = lines.index('(')
        second = lines.index(')')
        third = lines.index('(', first +1)
        fourth = lines.index(')', second + 1)
        field = lines[first:fourth+1]
        lines = lines.replace(field, '')
            
    while '"' in lines:
        first = lines.find('"')
        second = lines.find('"', first+1)
        field = lines[first:second+1]
        clean_field = field.replace(',',"")[1:-1]
        lines = lines.replace(field, clean_field)
    
    return lines



```

```python
with open(DATA) as file:
    for i in range(5):
        print(clean(file.readline()))
```

* Print each line (of the first 5, other than the header, and with the whitespace removed) followed by the indices of the two double quotes:
```
Argentina,ARG,"Feb 18, 2020",0
14 27
Argentina,ARG,"Feb 19, 2020",0
14 27
Argentina,ARG,"Feb 20, 2020",0
14 27
...
```

Tip: You can get the double quote character by enclosing it in single quotes (`'"'`).

Hint: Compare the `find` and `index` methods. Why would you choose one or the other?


* Next, print the line followed by the date string:

```
Argentina,ARG,"Feb 18, 2020",0
"Feb 18, 2020"
Argentina,ARG,"Feb 19, 2020",0
"Feb 19, 2020"
Argentina,ARG,"Feb 20, 2020",0
"Feb 20, 2020"
```

* Now do the same, except with the comma removed from the date string:
```Argentina,ARG,"Feb 18, 2020",0
Feb 18 2020
Argentina,ARG,"Feb 19, 2020",0
Feb 19 2020
Argentina,ARG,"Feb 20, 2020",0
Feb 20 2020    
```

```python

```

* Next, print the original lines with the old date field replaced by the cleaned date field:
```
Argentina,ARG,Feb 18 2020,0
Argentina,ARG,Feb 19 2020,0
Argentina,ARG,Feb 20 2020,0
```

#### Checked Solution [1 lab mark] 
Now that we have this working, we don't want this preprocessing step 'muddying' up our code, so let's put it in a separate function.

* Write a function `clean (data_row)` that takes as its argument a string, and:
  * strips any unnecessary whitespace characters from the ends
  * if it contains a string in double quotes, strips the quotes and the comma between them
  * if it doesn't contain any quotes (this will be the header line), it strips everything in the line from the space before the first parenthesis

If called with the following code:
```
with open(DATA,'r') as file:
    for i in range(5):
        print(clean(file.readline()))
```
the output should start like this:
```
Entity,Code,Date,Daily tests per thousand people
Argentina,ARG,Feb 18 2020,0
Argentina,ARG,Feb 19 2020,0
...
```

```python

```

```python

```

<!-- grade_id: clean, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal
DATA = "daily-tests-per-thousand-people-smoothed-7-day-20200729.csv"
with open(DATA,'r') as file:
    assert_equal(clean(file.readline()), "Entity,Code,Date,Daily tests per thousand people")
    assert_equal(clean(file.readline()), "Argentina,ARG,Feb 18 2020,0")
print("So far, so good on the practice (non-hidden) tests. Remember there will be additional tests applied.")

```

#### From a file to a list

Next, rather than printing the lines, store them all in a list (of lists).

* Define a variable `data_lists` as an empty list (). Write code that cleans and splits the (entire) input into lists, and appends them to `data_lists`.

For example, the following should print the first five rows as a list of lists:

```
print(data_lists[:5])

[['Entity', 'Code', 'Date', 'Daily tests per thousand people'], ['Argentina', 'ARG', 'Feb 18 2020', '0'], ['Argentina', 'ARG', 'Feb 19 2020', '0'], ['Argentina', 'ARG', 'Feb 20 2020', '0'], ['Argentina', 'ARG', 'Feb 21 2020', '0']]
```

How many entries (lines of data) are there in the file?

```python
data_lists = []
        
with open(DATA) as file:
    for line in file:

        lines = line.strip()
        
        
#         while ',,' in lines:
#             lines = lines.replace(',,',',')

        while '(' in lines:
            first = lines.index('(')
            second = lines.index(')')
            third = lines.index('(', first +1)
            fourth = lines.index(')', second + 1)
            field = lines[first:fourth+1]
            lines = lines.replace(field, '')
        

        while '"' in lines:
            first = lines.find('"')
            second = lines.find('"', first+1)
            field = lines[first:second+1]
            clean_field = field.replace(',',"")[1:-1]
            lines = lines.replace(field, clean_field)
        
        data_lists.append(lines.split(','))
data_lists[:5]
```

Output:
```
[['Entity', 'Code', 'Date', 'Daily tests per thousand people '],
 ['Argentina', 'ARG', 'Feb 18 2020', '0'],
 ['Argentina', 'ARG', 'Feb 19 2020', '0'],
 ['Argentina', 'ARG', 'Feb 20 2020', '0'],
 ['Argentina', 'ARG', 'Feb 21 2020', '0']]
```

```python
data_lists
```

Output:
```
[['Entity', 'Code', 'Date', 'Daily tests per thousand people '],
 ['Argentina', 'ARG', 'Feb 18 2020', '0'],
 ['Argentina', 'ARG', 'Feb 19 2020', '0'],
 ['Argentina', 'ARG', 'Feb 20 2020', '0'],
 ['Argentina', 'ARG', 'Feb 21 2020', '0'],
 ['Argentina', 'ARG', 'Feb 22 2020', '0'],
 ['Argentina', 'ARG', 'Feb 23 2020', '0'],
 ['Argentina', 'ARG', 'Feb 24 2020', '0'],
 ['Argentina', 'ARG', 'Feb 25 2020', '0'],
 ['Argentina', 'ARG', 'Feb 26 2020', '0'],
 ['Argentina', 'ARG', 'Feb 27 2020', '0'],
 ['Argentina', 'ARG', 'Feb 28 2020', '0'],
 ['Argentina', 'ARG', 'Feb 29 2020', '0'],
 ['Argentina', 'ARG', 'Mar 1 2020', '0'],
 ['Argentina', 'ARG', 'Mar 2 2020', '0'],
 ['Argentina', 'ARG', 'Mar 3 2020', '0'],
 ['Argentina', 'ARG', 'Mar 4 2020', '0'],
 ['Argentina', 'ARG', 'Mar 5 2020', '0'],
 ['Argentina', 'ARG', 'Mar 6 2020', '0'],
 ['Argentina', 'ARG', 'Mar 7 2020', '0'],
 ['Argentina', 'ARG', 'Mar 8 2020', '0'],
 ['Argentina', 'ARG', 'Mar 9 2020', '0'],
 ['Argentina', 'ARG', 'Mar 10 2020', '0']
```

#### Checking our cleaning so far

We now have a tidy list of lists, each with the four fields. Or do we?

With big data it may not be possible to manually look at every entry to see if we've accounted for every possibility. We should try to make our cleaning or preprocessing as general as possible so that we catch unexpected variations or bad data. 

In practice, we often have to make some assumptions about the data. However we should endeavour to test these.

In this case, we've assumed that the patterns we see at the start of the file continue through the file. So let's check that assumption.

* Write code that checks whether all your entries have 4 fields.

If not, why not? What have we missed?

Hint: Print out the first row (if there is one) where this is not true. Print out the number of that row, open the data file in CoCalc, and have a look at the data in that row. What do you find?


```python
# found = False
# count = 0
# for row in data_lists:
#     count = count + 1
#     if not found and len(row) !=4:
#         print(count,row)
#         found = True

has4 = []

for row in data_lists:
    has4.append(len(row) !=4)


has4.index(True)

```

Output:
```
ValueError                                Traceback (most recent call last)
<ipython-input-10-bfa5ca84d90f> in <module>
     13 
     14 
---> 15 has4.index(True)

ValueError: True is not in list
```

```python
has4 = []
for data in data_lists:
    has4.append(len(data)!=4)

first= has4.index(True)
data_lists[first]
```

Output:
```
ValueError                                Traceback (most recent call last)
<ipython-input-11-c5843f051181> in <module>
      3     has4.append(len(data)!=4)
      4 
----> 5 first= has4.index(True)
      6 data_lists[first]

ValueError: True is not in list
```

```python

```

#### General vs Specific

The variations you see in the data are not unusual - remember that this is the collated official government data, its not an 'exercise'. It is *exactly the kind of thing you would deal with as a working Data Scientist*.

We'll need to come back to some differences in the content of the data, but for now let's focus on the formatting. Or, more precisely, *transforming* the data from the format in which it is provided to a format that is suitable for our use.

* Alter your code to do more cleaning as necessary so that it is transformed into a list of lists, each with four fields.

You should try to make your code as **general as possible**. This means that, rather than just adjust for the specific case, think about *patterns*. 

For example, we have seen that the data format uses double quotes around fields containing the delimiter (a comma in this case). If the data is not corrupt (which is an assumption) therefore, we would expect the double quotes to always occur in pairs.
By focussing only on dates, we have been more *specific* than we need to be, so we may miss other cases. A more *general* solution will assume that the same pattern may occur in other fields.

Ensure you re-test your code after any changes you make to ensure it satisfies the requirements. (You might find it useful to write them all down.)

#### Type casting

Finally, the last field, tests per 1000 people, should be a `float`.

* Alter your code to change the tests ratio to a float.

Your first few lines should now look like this:
```
[['Entity', 'Code', 'Date', 'Daily tests per thousand people'], ['Argentina', 'ARG', 'Feb 18 2020', 0.0], ['Argentina', 'ARG', 'Feb 19 2020', 0.0], ['Argentina', 'ARG', 'Feb 20 2020', 0.0], ['Argentina', 'ARG', 'Feb 21 2020', 0.0]]
```

Again, we could make an assumption that the fourth string is always able to be converted to a float, but its possible that somewhere in the file that is not true. Later we will deal with this using *Exceptions*. For now, it is good practice to do some checks before attempting to cast. Before casting, check that:
* the fourth field is not an empty string
* the first character in the string is a number

Tip: You may find string functions like `isnumeric` useful.

```python
first = True
for line in data_lists:
    if first:
        first = False
    else:
        if line[3] != '' and line[3][0].isnumeric():
            x = float(line[3])
            line[3] = x
            
            
    
            
            
            
```

```python
data_lists
```

Output:
```
[['Entity', 'Code', 'Date', 'Daily tests per thousand people '],
 ['Argentina', 'ARG', 'Feb 18 2020', 0.0],
 ['Argentina', 'ARG', 'Feb 19 2020', 0.0],
 ['Argentina', 'ARG', 'Feb 20 2020', 0.0],
 ['Argentina', 'ARG', 'Feb 21 2020', 0.0],
 ['Argentina', 'ARG', 'Feb 22 2020', 0.0],
 ['Argentina', 'ARG', 'Feb 23 2020', 0.0],
 ['Argentina', 'ARG', 'Feb 24 2020', 0.0],
 ['Argentina', 'ARG', 'Feb 25 2020', 0.0],
 ['Argentina', 'ARG', 'Feb 26 2020', 0.0],
 ['Argentina', 'ARG', 'Feb 27 2020', 0.0],
 ['Argentina', 'ARG', 'Feb 28 2020', 0.0],
 ['Argentina', 'ARG', 'Feb 29 2020', 0.0],
 ['Argentina', 'ARG', 'Mar 1 2020', 0.0],
 ['Argentina', 'ARG', 'Mar 2 2020', 0.0],
 ['Argentina', 'ARG', 'Mar 3 2020', 0.0],
 ['Argentina', 'ARG', 'Mar 4 2020', 0.0],
 ['Argentina', 'ARG', 'Mar 5 2020', 0.0],
 ['Argentina', 'ARG', 'Mar 6 2020', 0.0],
 ['Argentina', 'ARG', 'Mar 7 2020', 0.0],
 ['Argentina', 'ARG', 'Mar 8 2020', 0.0],
 ['Argentina', 'ARG', 'Mar 9 2020', 0.0],
 ['Argentina', 'ARG', 'Mar 10 2020', 0.0]
```

```python

```

```python

```

```python
with open(DATA) as file:
    for line in file:
        print(line)
```

Output:
```
Entity,Code,Date,Daily tests per thousand people (7-day smoothed) (tests per thousand)

Argentina,ARG,"Feb 18, 2020",0

Argentina,ARG,"Feb 19, 2020",0

Argentina,ARG,"Feb 20, 2020",0

Argentina,ARG,"Feb 21, 2020",0

Argentina,ARG,"Feb 22, 2020",0

Argentina,ARG,"Feb 23, 2020",0

Argentina,ARG,"Feb 24, 2020",0

Argentina,ARG,"Feb 25, 2020",0

Argentina,ARG,"Feb 26, 2020",0

Argentina,ARG,"Feb 27, 2020",0

Argentina,ARG,"Feb 28, 2020",0

Argentina,ARG,"Feb 29, 2020",0

Argentina,ARG,"Mar 1, 2020",0

Argentina,ARG,"Mar 2, 2020",0

Argentina,ARG,"Mar 3, 2020",0

Argentina,ARG,"Mar 4, 2020",0

Argentina,ARG,"Mar 5, 2020",0

Argentina,ARG,"Mar 6, 2020",0

Argentina,ARG,"Mar 7, 2020",0

Argentina,ARG,"Mar 8, 2020",0

Argentina,ARG,"Mar 9, 2020",0

Argentina,ARG,"Mar 10, 2020",0

Argentina,ARG,"Mar 11, 2020",0

Argentina,ARG,"Mar 12, 2020",0

Argentina,ARG,"Mar 13, 2020",0.001

Argentina,ARG,"Mar 14, 2020",0.001

Argentina,ARG,"Mar 15, 2020",0.001

Argentina,ARG,"Mar 16, 2020",0.001

Argentina
```

```python

is_first_line = True
for line in data_lists:
    if is_first_line:
        is_first_line = False
    else:
        if line[3] != '' and line[3][0].isnumeric():
            x = float(line[3])
            line[3] = x
            
            
            
            
            
        
```

Output:
```
TypeError                                 Traceback (most recent call last)
<ipython-input-15-0104b6808e19> in <module>
      4         is_first_line = False
      5     else:
----> 6         if line[3] != '' and line[3][0].isnumeric():
      7             x = float(line[3])
      8             line[3] = x

TypeError: 'float' object is not subscriptable
```

```python
float_lists = []
```

Remember to run your own tests, not rely on my tests or my sample output. For example, I've included the first few lines of output (the file is too large to include them all) but as we've seen those lines may not be indicative of the file as a whole. For one thing, they are all cases where the daily test ratio is zero.

Again, remember that a Data Scientist is like a detective - always thinking about what we could possibly have missed, and testing for it.

#### Checked Solution [2 marks]

* Complete the function `get_cleaned_lists (filename)` so that it returns a list of lists, each containing the fields from the data file, with quotes, commas, and leading/trailing whitespace characters removed, any parenthesised text removed, and the daily tests ratio as a float.

Note: Your checked functions may call preceding functions that you have written - they do not have to all be in one long function. As always, however, you should ensure you validate your function with a "clean" kernel.

```python
def get_cleaned_lists(filename):
    data_lists = []
    with open(DATA) as file:
        for line in file:
            while '(' in line:
                first = line.index('(')
                last = line.index(')', first+1)
                field = line[first-1:last+1]
                line = line.replace(field,'')
            while '"' in line:
                first = line.index('"')
                last = line.index('"', first+1)
                field = line[first:last+1]
                clean_field = field.replace(',',"")[1:-1]
                line = line.replace(field, clean_field)
            data_lists.append(line.strip().split(','))
            
    is_first_line = True
    for line in data_lists:
        if is_first_line:
            is_first_line = False
        else:
            if line[3] != '' and line[3][0].isnumeric():
                x = float(line[3])
                line[3] = x
                
                
    return data_lists
```

<!-- grade_id: get_cleaned_lists-1, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal
DATA = "daily-tests-per-thousand-people-smoothed-7-day-20200729.csv"
assert_equal(get_cleaned_lists(DATA)[0],['Entity', 'Code', 'Date', 'Daily tests per thousand people'])
assert_equal(get_cleaned_lists(DATA)[5],['Argentina', 'ARG', 'Feb 22 2020', 0.0])
assert_equal(get_cleaned_lists(DATA)[40],['Argentina', 'ARG', 'Mar 28 2020', 0.008])
assert_equal(get_cleaned_lists(DATA)[180],['Argentina tests performed', '', 'Mar 10 2020', 0.0])
print("So far, so good. Remember there will be additional tests applied.")

```

<!-- grade_id: get_cleaned_lists_2, points: 1, grade: true, solution: false -->
```python
# For marking use only.

```

- do not include the first line
- take care of empty space
- take care of first digit being number
- change to float


```python
first_line = True
for line in data_lists:
    if first_line:
        first_line = False
    else:
        if line == '' and line[3][0].isnumeric():
            x = float(line[3])
            line[3] = x

            
```

## Data Presentation  - Using Dictionaries

*For this part it is left to you to break down the task into subtasks and test them as you go.*

We want to be able to access the daily test ratio for a country on a given date without using loops.

Dictionaries provide *much* faster access to data by *hashing* the dictionary keys.

* Store the daily tests data in a dictionary of dictionaries. The outer dictionary should use the countries as keys. The inner dictionary should use the dates as keys.

For example, if my outer dictionary is called `country_dict`, then evaluating `country_dict["Australia"]` should return:
    
```
{'Mar 29 2020': 0.382,
 'Mar 30 2020': 0.396,
 'Mar 31 2020': 0.41,
 'Apr 1 2020': 0.425,
 'Apr 2 2020': 0.439,
 'Apr 3 2020': 0.453,
 'Apr 4 2020': 0.467,
 ...
```
and `country_dict["Australia"]["Jul 1 2020"]` should return:
```
1.824
```






```python



```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python
country_dict = {}
for line in data_lists[1:]:
    country, date, test_ratio = line[0],line[2],line[3]
    if country not in country_dict:
        country_dict[country] = {}
    country_dict[country][date] = test_ratio

country_dict
```

Output:
```
{'Argentina': {'Feb 18 2020': 0.0,
  'Feb 19 2020': 0.0,
  'Feb 20 2020': 0.0,
  'Feb 21 2020': 0.0,
  'Feb 22 2020': 0.0,
  'Feb 23 2020': 0.0,
  'Feb 24 2020': 0.0,
  'Feb 25 2020': 0.0,
  'Feb 26 2020': 0.0,
  'Feb 27 2020': 0.0,
  'Feb 28 2020': 0.0,
  'Feb 29 2020': 0.0,
  'Mar 1 2020': 0.0,
  'Mar 2 2020': 0.0,
  'Mar 3 2020': 0.0,
  'Mar 4 2020': 0.0,
  'Mar 5 2020': 0.0,
  'Mar 6 2020': 0.0,
  'Mar 7 2020': 0.0,
  'Mar 8 2020': 0.0,
  'Mar 9 2020': 0.0,
  'Mar 10 2020': 0.0,
  'Mar 11 2020': 0.0,
  'Mar 12 2020': 0.0,
  'Mar 13 2020': 0.001,
  'Mar 14 2020': 0.001,
  'Mar 15 2020': 0.001,
  'Mar 16 2020': 0.001,
  'Mar 17 2020': 0.002,
  'Mar 18 2020': 0.002,
  'Mar 19 2020': 0.002,
  'Mar 20 2020': 0.003,
  'Mar 21 2020': 0.004,
  'Mar 22 2020': 0.004,
  'Mar 23 2020': 0.005,
  'Mar 24 2020': 0.005,
  'Mar 25 2020': 0.006,
  'Mar 26 2020': 0.007,
  'Mar 27 2020': 0.008,
  'Mar 28 2020': 0.008,
  'Mar 29 2020': 0.009,
  'Mar 30 2020': 0.01,
  'Mar 31 2020': 0.011,
  'Apr 1 2020
```

```python
country_dict["Australia"]["Jul 1 2020"]
```

Output:
```
1.824
```

```python

```

```python

```

### Presenting rankings in a table

* Finally, write a function `print_ranking(date)` that takes a date (as a string) and prints a table of testing results for that date, ranked from highest testing rate to the lowest.

So, for example, `print_ranking("Jul 1 2020")` will start as follows:
    
```
Testing results for Jul 1 2020
8.032 	 Luxembourg
5.653 	 United Arab Emirates
5.055 	 Bahrain
...
```

```python
country_ranking_ratio = []

date = "Jul 1 2020"

for country in country_dict:
    if date in country_dict[country]:
        country_ranking_ratio.append((country_dict[country][date], country))
            
    

```

```python
sorted(country_ranking_ratio)
        
        

```

Output:
```
[(0.007, 'Taiwan'),
 (0.009, 'Democratic Republic of Congo'),
 (0.013, 'Nigeria'),
 (0.027, 'Zimbabwe'),
 (0.028, 'Myanmar'),
 (0.028, 'Thailand'),
 (0.03, 'Thailand tests performed'),
 (0.033, 'Ethiopia'),
 (0.033, 'Tunisia'),
 (0.041, 'Indonesia'),
 (0.043, 'Japan'),
 (0.045, 'Togo'),
 (0.051, "Cote d'Ivoire"),
 (0.053, 'Uganda'),
 (0.058, 'Kenya'),
 (0.06, 'Japan tests performed'),
 (0.063, 'Senegal'),
 (0.075, 'Fiji'),
 (0.087, 'Mexico'),
 (0.095, 'Ghana'),
 (0.1, 'Pakistan'),
 (0.106, 'Ecuador'),
 (0.107, 'Bangladesh'),
 (0.127, 'Peru'),
 (0.141, 'Philippines'),
 (0.145, 'Costa Rica'),
 (0.153, 'Bolivia'),
 (0.153, 'India'),
 (0.163, 'Argentina'),
 (0.196, 'Nepal'),
 (0.198, 'Paraguay'),
 (0.199, 'Cuba'),
 (0.202, 'Argentina tests performed'),
 (0.207, 'Slovakia'),
 (0.216, 'South Korea'),
 (0.234, 'Croatia'),
 (0.236, 'Ukraine'),
 (0.242, 'Hungary'),
 (0.29, 'Uruguay'),
 (0.291, 'Rwanda'),
 (0.307, 'Malaysia'),
 (0.324, 'Iran'),
 (0.335, 'Greece'),
 (0.357, 'Colombia'),
 (0.361, 
```

```python

```

```python

```

```python

```

```python
def print_ranking(date):
    
    country_test_ratio = []
    for country in country_dict:
        if date in country_dict[country]:
            country_test_ratio.append((country_dict[country][date],country))
    
    return sorted(country_test_ratio, reverse= True)
    
```

```python

```

```python
print_ranking("Jul 1 2020")
```

Output:
```
[(8.032, 'Luxembourg'),
 (5.653, 'United Arab Emirates'),
 (5.055, 'Bahrain'),
 (2.61, 'Denmark'),
 (2.316, 'United States tests performed '),
 (2.038, 'Singapore samples tested'),
 (2.01, 'Russia'),
 (1.97, 'Israel'),
 (1.869, 'United States'),
 (1.824, 'Australia'),
 (1.807, 'Malta'),
 (1.754, 'Belarus'),
 (1.689, 'Maldives'),
 (1.392, 'United Kingdom'),
 (1.355, 'Qatar'),
 (1.28, 'Lithuania'),
 (1.264, 'Portugal'),
 (1.208, 'Saudi Arabia'),
 (1.135, 'Switzerland'),
 (1.093, 'New Zealand'),
 (1.087, 'Canada'),
 (1.078, 'Sweden people tested'),
 (1.066, 'Serbia'),
 (1.026, 'Kazakhstan'),
 (1.024, 'Singapore'),
 (0.949, 'Iceland'),
 (0.864, 'Kuwait'),
 (0.841, 'Chile'),
 (0.838, 'Latvia'),
 (0.823, 'Belgium'),
 (0.817, 'Germany'),
 (0.8, 'Italy'),
 (0.749, 'Ireland'),
 (0.671, 'Austria'),
 (0.657, 'Oman'),
 (0.641, 'Panama'),
 (0.602, 'South Africa'),
 (0.597, 'France'),
 (0.594, 'Turkey'),
 (0.594, 'Norway'),
 (0.554, 'Poland'),
 (0.553, 'Netherlands'),
 (0.535, 'Spain'),
 (0.529, 'Ro
```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

### Checking the news

The Whitehouse statement cited in the introduction, above, was made on April 28th.

What do you make of President Trump's statements *from a per capita perspective*?

How does that compare with more recently?


In a May 11 [Rose Garden briefing](https://www.whitehouse.gov/briefings-statements/remarks-president-trump-press-briefing-covid-19-testing/) President Trump stated:

> We’re testing more people per capita than South Korea, the United Kingdom, France, Japan, Sweden, Finland, and many other countries — and, in some cases, combined.

The BBC's May 15th article [Coronavirus: President Trump’s testing claims fact-checked](https://www.bbc.com/news/world-us-canada-52493073) "fact-checks" this claim (Claim One).

* Modify your function to the signature `print_ranking(date, countries=[])` so that:
  * if `countries` is omitted, it still prints the table for *all* countries reporting on that day
  * if a list of countries is passed to the function, then it only prints the table for those countries

Print the table for the US and those six countries on 11th May.

Is the BBC's fact check for the 11th May borne out by these data?

Can you think of a possible reason for these discrepancies? (Hint: Should "we're testing" be interpreted as a rate or as a cumulative total?)

```python
def print_ranking(date, countries=[]):
    countries_ratio = []

    

    countries_ratio.sort(reverse = True)

    return countries_ratio

```

```python
print_ranking("Jul 1 2020", countries=['Australia', 'Argentina'])
```

Output:
```
[]
```

On 22nd June Newsweek, in [Why Trump Is Both Right and Wrong About U.S. Coronavirus Testing Numbers](https://www.newsweek.com/why-trump-right-wrong-about-us-coronavirus-testing-1512472), compares the US with Russia, Spain, Germany and Portugal on cumulative per capita figures. 

How does this compare with the picture you get for the daily rate at this date?

Congratulations - you can now get a job as a fact checking journalist!

&copy; Cara MacNish, Univeristy of Western Australia

```python

```

```python

```

```python

```

```python

```

```python

```