<span style="font-size:x-large;">**Lab 1**</span>

# Diversity in Study: What (Else) Are We All Studying?

----


## Data Science Across the Disciplines

Data Science has application right across the disciplines. In this course we'll see case studies ranging from the Environment to Economics to Health.

Before we start this journey, it's interesting to see what diverse areas the group is studying. As we said in the Welcome lecture, there are many students in the unit with many different study backgrounds and focusses.

To get a feel for this, we'll look at the range of enrolment patterns for this year.

Note that some randomisation has been applied to the enrolment patterns to ensure privacy and anonymity. However the disciplines are preserved.

#### Reference links

In this lab we will practise reading in data from a file, cleaning the data, and observing features of the data.

We'll be using standard python. The latest documentation will be useful, and it is suggested you bookmark it for future use.

* _The Python Language Reference_, https://docs.python.org/3/reference/
* _The Python Standard Library_, https://docs.python.org/3/library/
* _The Python Tutorial_, https://docs.python.org/3/tutorial/

## Data Acquisition

In future labs we'll be hunting down data as part of the work. In this case however the data will be provided in a text file.

* Open the directory from the folder icon. Click on the file `unit-patterns.txt` to inspect the file.

The file contains all the different study plans of enrolled students (at the time of writing) with some obfuscation for privacy and anonymity.

This is a *tab-separated values (tsv)* file with two *fields*.

The first field contains a list of units studied (separated by plus signs), the second (after the tab character) contains a number and a percentage.

Note that in this lab we will refer to an instance of a unit name combined with a teaching period as a 'unit' - so for example, CITS1401-1 and CITS1401-2 would be treated as two distinct units.

---

Let's start with a constant to identify the file. By convention we'll use uppercase for constants. Run the following cell to make the assignment.

```python
DATAFILE = "unit-patterns.txt"
```

```python

```

As we saw in the 'Getting Started' lab, we can check to see that the assignment was made just by typing the variable name as the last thing in the cell and running the cell. Try this now.

#### Reading the file

* Use the built-in function `open()` and the text file IO method `readline()` to read in and print out the first 5 lines of the file. 

Your output should start like this:
```
 CITS2224-1 + CITS2023-2 + CITS3221-2 + CITS3224-1 + CITS3225-1 	 1 (2.0%) 

 CITS2023-2 + ENSC3221-2 + ENSC3224-1 + ENSC3225-1 + ENSC3227-2 + ENSC3210-1 + ENSC3221-1 + MECH3023-2 	 1 (2.0%) 

 CITS1021-2 + CITS2021-1 + CITS2023-2 + ENSC2224-1 + ENSC3211-2 + ENSC3213-2 + MATH2221-1 + STAT1524-1 	 1 (2.0%) 
...
```

Tip: If you're a little rusty on reading files, you may find this tutorial useful: [7.2. Reading and Writing Files](https://docs.python.org/3/tutorial/inputoutput.html?highlight=readline#reading-and-writing-files).

Also the Case Study video "CITS2402 FakeNews 2" provides a couple of examples of working with a tsv file.



```python

with open(DATAFILE) as file:
    for i in range(5):
        line = file.readline()

```

* Read the file again, but this time, as you read the lines from the file, save them in a list called `pattern-strings`.

Printing the first 5 lines of your list should give:
```
[' CITS2224-1 + CITS2023-2 + CITS3221-2 + CITS3224-1 + CITS3225-1 \t 1 (2.0%) \n', ' CITS2023-2 + ENSC3221-2 + ENSC3224-1 + ENSC3225-1 + ENSC3227-2 + ENSC3210-1 + ENSC3221-1 + MECH3023-2 \t 1 (2.0%) \n', ' CITS1021-2 + CITS2021-1 + CITS2023-2 + ENSC2224-1 + ENSC3211-2 + ENSC3213-2 + MATH2221-1 + STAT1524-1 \t 1 (2.0%) \n', ' CITS2023-2 + CITS3220-2 + CITS3223-2 + CITS3021-1 + CITS3025-1 + MATH3224-1 + MATH3221-1 + STAT3264-1 + STAT3260-2 \t 1 (2.0%) \n', ' BIOC3221-1 + BIOC3224-1 + BIOC3225-2 + BIOC3225-2 + CITS1023-2 + CITS2021-1 + CITS2023-2 + MATH1211-1 \t 1 (2.0%) \n']
```

Tip: Rather than using `readline()`, use an *iterator* on the file directly to get each line in turn.

```python
pattern_strings = []

with open(DATAFILE) as file:
    for lines in file:
        pattern_strings.append(lines)
pattern_strings[:5]
```

Output:
```
[' CITS2224-1 + CITS2023-2 + CITS3221-2 + CITS3224-1 + CITS3225-1 \t 1 (2.0%) \n',
 ' CITS2023-2 + ENSC3221-2 + ENSC3224-1 + ENSC3225-1 + ENSC3227-2 + ENSC3210-1 + ENSC3221-1 + MECH3023-2 \t 1 (2.0%) \n',
 ' CITS1021-2 + CITS2021-1 + CITS2023-2 + ENSC2224-1 + ENSC3211-2 + ENSC3213-2 + MATH2221-1 + STAT1524-1 \t 1 (2.0%) \n',
 ' CITS2023-2 + CITS3220-2 + CITS3223-2 + CITS3021-1 + CITS3025-1 + MATH3224-1 + MATH3221-1 + STAT3264-1 + STAT3260-2 \t 1 (2.0%) \n',
 ' BIOC3221-1 + BIOC3224-1 + BIOC3225-2 + BIOC3225-2 + CITS1023-2 + CITS2021-1 + CITS2023-2 + MATH1211-1 \t 1 (2.0%) \n']
```

## Data Cleaning

We are only interested in the patterns to the left of the tab.

* Using the string `partition()` method, modify your code so that `pattern_strings` contains only the unit patterns.

The documentation for [String Methods](https://docs.python.org/3/library/stdtypes.html#string-methods) may be helpful.



Printing the first 5 lines of your list (`print(pattern_strings[:5])`) should now give:
```
[' CITS2224-1 + CITS2023-2 + CITS3221-2 + CITS3224-1 + CITS3225-1 ', ' CITS2023-2 + ENSC3221-2 + ENSC3224-1 + ENSC3225-1 + ENSC3227-2 + ENSC3210-1 + ENSC3221-1 + MECH3023-2 ', ' CITS1021-2 + CITS2021-1 + CITS2023-2 + ENSC2224-1 + ENSC3211-2 + ENSC3213-2 + MATH2221-1 + STAT1524-1 ', ' CITS2023-2 + CITS3220-2 + CITS3223-2 + CITS3021-1 + CITS3025-1 + MATH3224-1 + MATH3221-1 + STAT3264-1 + STAT3260-2 ', ' BIOC3221-1 + BIOC3224-1 + BIOC3225-2 + BIOC3225-2 + CITS1023-2 + CITS2021-1 + CITS2023-2 + MATH1211-1 ']
```

* How many students were enrolled at the time the list was taken?

```python
pattern_strings = []

with open(DATAFILE) as file:
    for lines in file:
        pattern_strings.append(lines.partition("\t")[0])
print(len(pattern_strings))
pattern_strings
```

Output:
```
261
[' CITS2224-1 + CITS2023-2 + CITS3221-2 + CITS3224-1 + CITS3225-1 ',
 ' CITS2023-2 + ENSC3221-2 + ENSC3224-1 + ENSC3225-1 + ENSC3227-2 + ENSC3210-1 + ENSC3221-1 + MECH3023-2 ',
 ' CITS1021-2 + CITS2021-1 + CITS2023-2 + ENSC2224-1 + ENSC3211-2 + ENSC3213-2 + MATH2221-1 + STAT1524-1 ',
 ' CITS2023-2 + CITS3220-2 + CITS3223-2 + CITS3021-1 + CITS3025-1 + MATH3224-1 + MATH3221-1 + STAT3264-1 + STAT3260-2 ',
 ' BIOC3221-1 + BIOC3224-1 + BIOC3225-2 + BIOC3225-2 + CITS1023-2 + CITS2021-1 + CITS2023-2 + MATH1211-1 ',
 ' CITS2021-1 + CITS2023-2 + GRMN2025-1 + GRMN2020-2 + MATH1214-1 + MATH2521-2 + PHYS2221-1 + PHYS2223-2 ',
 ' CITS1221-1 + CITS1021-1 + CITS1023-2 + CITS2023-2 + MGMT1135-1 + MGMT1136-2 + STAT1024-1 + STAT2023-2 ',
 ' CITS1221-2 + CITS2023-2 + CITS3223-2 + CITS3021-1 + CITS3025-1 + CLAN3224-1 + STAT2021-1 + STAT3260-2 ',
 ' CITS2023-2 + CITS3025-1 + ECON1121-2 + FINA2229-1 + FINA2224-1 + MGMT1135-1 + MKTG1225-2 + PHAR1121-2 ',
 ' CITS2023-2 + CITS3220-2 + CITS3223-2 + STAT2023
```

Rather than the units in each line being part of a string, it is more useful to have them in a list. This will enable us to iterate through the list later.

* Use the `split` method to split each line into a list of units. Store the result in a variable `pattern_lists`.

This should result in a list of lists. Your output should start like this:
```
[[' CITS2224-1 ', ' CITS2023-2 ', ' CITS3221-2 ', ' CITS3224-1 ', ' CITS3225-1 '], [' CITS2023-2 ', ' ENSC3221-2 ', ' ENSC3224-1 ', ' ENSC3225-1 ', ' ENSC3227-2 ', ' ENSC3210-1 ', ' ENSC3221-1 ', ' MECH3023-2 '], [' CITS1021-2 ', ' CITS2021-1 ', ' CITS2023-2 ', ' ENSC2224-1 ', ' ENSC3211-2 ', ' ENSC3213-2 ', ' MATH2221-1 ', ' STAT1524-1 '], [' CITS2023-2 ', ' CITS3220-2 ', ' CITS3223-2 ', ' CITS3021-1 ', ' CITS3025-1 ', ' MATH3224-1 ', ' MATH3221-1 ', ' STAT3264-1 ', ' STAT3260-2 '], [' BIOC3221-1 ', ' BIOC3224-1 ', ' BIOC3225-2 ', ' BIOC3225-2 ', ' CITS1023-2 ', ' CITS2021-1 ', ' CITS2023-2 ', ' MATH1211-1 ']]
```



```python
print(type(pattern_strings[0]))


pattern_strings = []
with open(DATAFILE) as file:
    for lines in file:
        pattern_strings.append(lines.partition("\t")[0].strip().split(' + '))

print(pattern_strings[:5])


```

Output:
```
<class 'str'>
[['CITS2224-1', 'CITS2023-2', 'CITS3221-2', 'CITS3224-1', 'CITS3225-1'], ['CITS2023-2', 'ENSC3221-2', 'ENSC3224-1', 'ENSC3225-1', 'ENSC3227-2', 'ENSC3210-1', 'ENSC3221-1', 'MECH3023-2'], ['CITS1021-2', 'CITS2021-1', 'CITS2023-2', 'ENSC2224-1', 'ENSC3211-2', 'ENSC3213-2', 'MATH2221-1', 'STAT1524-1'], ['CITS2023-2', 'CITS3220-2', 'CITS3223-2', 'CITS3021-1', 'CITS3025-1', 'MATH3224-1', 'MATH3221-1', 'STAT3264-1', 'STAT3260-2'], ['BIOC3221-1', 'BIOC3224-1', 'BIOC3225-2', 'BIOC3225-2', 'CITS1023-2', 'CITS2021-1', 'CITS2023-2', 'MATH1211-1']]
```

Finally, notice that we've been left with unnecessary whitespace around each unit name.

* Modify your code to strip off the extra whitespace.

Your first 5 patterns should now look like this:
    
```
[['CITS2224-1', 'CITS2023-2', 'CITS3221-2', 'CITS3224-1', 'CITS3225-1'], ['CITS2023-2', 'ENSC3221-2', 'ENSC3224-1', 'ENSC3225-1', 'ENSC3227-2', 'ENSC3210-1', 'ENSC3221-1', 'MECH3023-2'], ['CITS1021-2', 'CITS2021-1', 'CITS2023-2', 'ENSC2224-1', 'ENSC3211-2', 'ENSC3213-2', 'MATH2221-1', 'STAT1524-1'], ['CITS2023-2', 'CITS3220-2', 'CITS3223-2', 'CITS3021-1', 'CITS3025-1', 'MATH3224-1', 'MATH3221-1', 'STAT3264-1', 'STAT3260-2'], ['BIOC3221-1', 'BIOC3224-1', 'BIOC3225-2', 'BIOC3225-2', 'CITS1023-2', 'CITS2021-1', 'CITS2023-2', 'MATH1211-1']]
```

##### Putting it all together

Congratulations, you have now 'wrangled' your data from its 'raw' format in the file, into a very useable form.

* To consolidate this section, turn your code into a function `get_patterns(filename)` that takes a filename, and returns a list of lists containing all the patterns in the file.

`print(get_patterns(DATAFILE)[:5])` should now print the same output as above.

```python
def get_patterns(filename):
    
    pattern_strings = []
    with open(filename) as file:
        for lines in file:
            pattern_strings.append(lines.partition("\t")[0].strip().split(' + '))
    return pattern_strings

```

```python
print(get_patterns(DATAFILE)[:5])
```

Output:
```
[['CITS2224-1', 'CITS2023-2', 'CITS3221-2', 'CITS3224-1', 'CITS3225-1'], ['CITS2023-2', 'ENSC3221-2', 'ENSC3224-1', 'ENSC3225-1', 'ENSC3227-2', 'ENSC3210-1', 'ENSC3221-1', 'MECH3023-2'], ['CITS1021-2', 'CITS2021-1', 'CITS2023-2', 'ENSC2224-1', 'ENSC3211-2', 'ENSC3213-2', 'MATH2221-1', 'STAT1524-1'], ['CITS2023-2', 'CITS3220-2', 'CITS3223-2', 'CITS3021-1', 'CITS3025-1', 'MATH3224-1', 'MATH3221-1', 'STAT3264-1', 'STAT3260-2'], ['BIOC3221-1', 'BIOC3224-1', 'BIOC3225-2', 'BIOC3225-2', 'CITS1023-2', 'CITS2021-1', 'CITS2023-2', 'MATH1211-1']]
```

## Checked Questions and Answers

Many of the labs this semester will have one or two checked (or 'autograded') answers. These will contribute to the (35%) practical marks for the semester.

The checked answers will generally have two parts, or two cells. 

##### Answer cells

The first cell is where you enter your code to be checked. This may contains a "stub" - that is, a part of the code to be completed. For example, it may be a function definition.

##### Test cells

The second cell is where the code to check the answer will go and is not editable. It may be blank, but will usually contain some example tests that you can use to test your own code before the answers are collected. There will be additional tests applied that are not shown.

##### Collection and deadlines

The notebook will be collected and marked, by running the tests, after the lab deadline. (As the same notebook will be collected, you may wish to take a backup copy of the notebook before starting work. Alternatively you can use the TimeTravel feature to revert the notebook if you need to start again.)

Everyone will automatically receive a one week "extension" from the week of the lab to finish off the work and to account for any unforeseen circumstances.

The deadline for credit will therefore be **11:59pm on the Friday of the week after the lab is released**.



#### Checked Solution [1 mark]

* Copy your solution for `get_patterns(filename)` to complete the following stub.

(We wouldn't normally enter the function a second time, it is only because we are introducing the graded answer for the first time.)


Before a final test of your code, it is worth restarting the kernel from the "Kernel" drop-down menu, and re-running your code. This will clean up any leftover 'junk' in the memory from running earlier versions of bits of code.

You may find, for example, that after you fixed or changed a piece of code, you forgot to change a variable name, but the code still worked because an old value for the variable was still in memory.

* Restart the kernel and test your code again to make sure it is doing everything correctly.


#### Running the grading tests

The test cells will often include some 'non-hidden' tests that you can run to see if your code passes. There will also be some 'hidden' tests that are only run once the code is collected after the deadline.

If your code produces the correct answers it will simply run successfully.

If your code does not pass the non-hidden tests, you will normally get an error message. For example, if I only stripped the whitespace from the left side and not both sides of the unit names, I might get an error message like this:

```
---------------------------------------------------------------------------
AssertionError                            Traceback (most recent call last)
<ipython-input-82-f1838aea55b1> in <module>
      1 from nose.tools import assert_equal
      2 DATA = "unit-patterns.txt"
----> 3 assert_equal(get_patterns(DATA)[0],['CITS2224-1', 'CITS2023-2', 'CITS3221-2', 'CITS3224-1', 'CITS3225-1'])
      4 assert_equal(get_patterns(DATA)[5],['CITS2021-1', 'CITS2023-2', 'GRMN2025-1', 'GRMN2020-2', 'MATH1214-1', 'MATH2521-2', 'PHYS2221-1', 'PHYS2223-2'])
...
668     def fail(self, msg=None):
    669         """Fail immediately, with the given message."""
--> 670         raise self.failureException(msg)
    671 
    672     def assertFalse(self, expr, msg=None):
AssertionError: Lists differ: ['CITS2224-1 ', 'CITS2023-2 ', 'CITS3221-2 ', 'CITS3224-1 ', 'CITS3225-1 '] != ['CITS2224-1', 'CITS2023-2', 'CITS3221-2', 'CITS3224-1', 'CITS3225-1']

First differing element 0:
'CITS2224-1 '
'CITS2224-1'

- ['CITS2224-1 ', 'CITS2023-2 ', 'CITS3221-2 ', 'CITS3224-1 ', 'CITS3225-1 ']
?             -              -              -              -              -

+ ['CITS2224-1', 'CITS2023-2', 'CITS3221-2', 'CITS3224-1', 'CITS3225-1']
```

The error messages may vary in their usefulness - they are not intended as a debugging tool, but often they will give some pointer as to what went wrong. In this case it is fairly clear that there is whitespace that shouldn't be there.

---

You can then correct your code and try again. 

The following test is worth 1 prac mark. Give it a try. You can run the tests as may times as you wish before the deadline.

<!-- grade_id: get_patterns, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal
DATA = "unit-patterns.txt"
assert_equal(get_patterns(DATA)[0],['CITS2224-1', 'CITS2023-2', 'CITS3221-2', 'CITS3224-1', 'CITS3225-1'])
assert_equal(get_patterns(DATA)[5],['CITS2021-1', 'CITS2023-2', 'GRMN2025-1', 'GRMN2020-2', 'MATH1214-1', 'MATH2521-2', 'PHYS2221-1', 'PHYS2223-2'])

```

## Data Inspection and Interpretation

Now let's have a closer look at the data. The file contains all the different study plans of enrolled students. Let's start by working out how many *different* or *unique* plans there are.

To make it easier to see what is going on, start by defining a better print function.

* Write a function `print_patterns (patterns, numlines)` that prints the first `numlines` patterns each on a new line.

For example:

print_patterns(get_patterns(DATA), 8)

```
['CITS2224-1', 'CITS2023-2', 'CITS3221-2', 'CITS3224-1', 'CITS3225-1']
['CITS2023-2', 'ENSC3221-2', 'ENSC3224-1', 'ENSC3225-1', 'ENSC3227-2', 'ENSC3210-1', 'ENSC3221-1', 'MECH3023-2']
['CITS1021-2', 'CITS2021-1', 'CITS2023-2', 'ENSC2224-1', 'ENSC3211-2', 'ENSC3213-2', 'MATH2221-1', 'STAT1524-1']
['CITS2023-2', 'CITS3220-2', 'CITS3223-2', 'CITS3021-1', 'CITS3025-1', 'MATH3224-1', 'MATH3221-1', 'STAT3264-1', 'STAT3260-2']
['BIOC3221-1', 'BIOC3224-1', 'BIOC3225-2', 'BIOC3225-2', 'CITS1023-2', 'CITS2021-1', 'CITS2023-2', 'MATH1211-1']
['CITS2021-1', 'CITS2023-2', 'GRMN2025-1', 'GRMN2020-2', 'MATH1214-1', 'MATH2521-2', 'PHYS2221-1', 'PHYS2223-2']
['CITS1221-1', 'CITS1021-1', 'CITS1023-2', 'CITS2023-2', 'MGMT1135-1', 'MGMT1136-2', 'STAT1024-1', 'STAT2023-2']
['CITS1221-2', 'CITS2023-2', 'CITS3223-2', 'CITS3021-1', 'CITS3025-1', 'CLAN3224-1', 'STAT2021-1', 'STAT3260-2']
```

```python
def print_patterns(patterns, numlines):
    for line in patterns[:numlines]:
        print(line)

    
```

```python
print_patterns(get_patterns(DATAFILE),5)
```

* Write a function that takes a pattern list, and returns the list in which:
  * the units in each list (pattern) are sorted alpha-numerically
  * the lists (patterns) themselves are sorted alpha-numerically
    

The following test code:
```
print_patterns(get_patterns(DATA),3)
print()
print_patterns(sort_patterns(get_patterns(DATA)),3)
```

should produce the following output:
    
```
['CITS2224-1', 'CITS2023-2', 'CITS3221-2', 'CITS3224-1', 'CITS3225-1']
['CITS2023-2', 'ENSC3221-2', 'ENSC3224-1', 'ENSC3225-1', 'ENSC3227-2', 'ENSC3210-1', 'ENSC3221-1', 'MECH3023-2']
['CITS1021-2', 'CITS2021-1', 'CITS2023-2', 'ENSC2224-1', 'ENSC3211-2', 'ENSC3213-2', 'MATH2221-1', 'STAT1524-1']

['AACE1223-2AA', 'ACCT1121-1', 'CARS1223-2AA', 'CITS1021-1', 'CITS2023-2', 'ECON1121-1', 'FINA1221-2', 'INDG1223-2AA', 'MGMT1135-2', 'MKTG1225-2', 'STAT1524-1']
['AACE1223-2AA', 'CARS1223-2AA', 'CITS1021-1', 'CITS1023-2', 'CITS2023-2', 'INDG1223-2AA', 'LAWS1113-2', 'PSYC1121-1', 'PSYC1123-2', 'SCIE1121-1', 'STAT1024-1']
['AACE1223-2AA', 'CARS1223-2AA', 'CITS1021-2', 'CITS1023-2', 'COMM1223-2', 'INDG1223-2AA', 'STAT1023-2']```
```

_Caution: If you are using python's built-in sorting methods, it is highly recommended that you use `sorted` to return a new sorted list, rather than `sort` to sort the list in-place. Methods that alter a data structure 'in-place' can lead to confusing behaviour when debugging._

```python
def sort_patterns(patterns):
    out_p = []
    for outer_pattern in patterns:
        out_p.append(sorted(outer_pattern))
    inner_pattern = sorted(out_p)
    return inner_pattern
    
```

```python
sorted_patterns = print_patterns(sort_patterns(get_patterns(DATAFILE)),len(pattern_strings))
```

* Write a function `remove_duplicates (patterns)` that returns a (sorted) list of patterns with any duplicate patterns removed.

How many *unique* study patterns are there in the class?

```python
unique_patterns = []
for unit in pattern_strings:
    if unit not in box:
        box.append(unit)
```

```python

def remove_duplicates(patterns):
    sorted_patterns = sort_patterns(patterns)
    new_patterns = [sorted_patterns[0]]
    for pattern in sorted_patterns[1:]:
        if pattern != new_patterns[:-1]:
            new_patterns.append(pattern)
    return new_patterns
        
```

```python
def remove_duplicates(patterns):
    unique_patterns = []
    for unit in patterns:
        if unit not in unit:
            unique_patterns.append(unit)
    return unique_patterns
```

```python
def get_unique_patterns(patterns):
    unique_patterns = remove_duplicates(sort_patterns(patterns))
    num_patterns = len(patterns)
    
    return (unique_patterns, num_patterns)

```

```python
(unique_patterns, num_patterns) = get_unique_patterns(get_patterns(DATAFILE))
unique_patterns
```

```python

```

Print the percentage (to one decimal place) of the total number of study plans in the class that are different.

You should find that we are indeed a diverse group!

#### Checked Solution [1 mark]

* Write a function `get_unique_patterns (patterns)` that returns a pair `(unique_patterns, num_patterns)` where:
  * `unique_patterns` contains a sorted list of sorted patterns,  with any duplicate patterns removed (as above)
  * `num_patterns` contains the number of unique patterns

<!-- grade_id: get_unique_patterns, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal
DATA = "unit-patterns.txt"
(patterns, num_patterns) = get_unique_patterns(get_patterns(DATA))
assert_equal(patterns[0],['AACE1223-2AA', 'ACCT1121-1', 'CARS1223-2AA', 'CITS1021-1', 'CITS2023-2', 'ECON1121-1', 'FINA1221-2', 'INDG1223-2AA', 'MGMT1135-2', 'MKTG1225-2', 'STAT1524-1'])
assert_equal(patterns[5][0],'AACE1224-ACE-2')
assert_equal(patterns[10][-1],'STAT1523-2')

```

## Examining the range of disciplines

For this part, it is left to you to think about the best way to break the problems down. You won't want to write new code for each different discipline, so think about how you can use structured programming to efficiently answer the questions.

* What proportion of study patterns contain: 
  * a maths unit?
  * a stats unit?
  * a business unit?
  * a medicine unit?
  * a law unit?
  * a psychology unit?
  * a music unit?
  * a philosophy unit?
  * a service learning unit?


* How many different disciplines are studied altogether (for the purposes of this question assume the first four characters of the unit code represent the discipline.)

```python
DATA = "unit-patterns.txt"
patterns = get_patterns(DATA)
print(patterns[:5])
```

# find proportion of each unit. But how?
### Steps of finding proportion

#### A: Find students study Math

- find how many students study at least two math clas
- find how many student study one math class
- Calculate proportion of students study one class,or two classes. 





```python
math_one = 0
for pattern in pattern_strings:
    count = 0
    for unit in pattern:
        if unit.startswith('MATH'):
            count+=1
    #print(pattern)
    if count >= 1:
        math_one+=1
print(math_one)
        
proportion = math_one/len(pattern_strings)*100



```

## Proportion of units

```python
def find_proportion(subject,data):
    students = 0
    for pattern in data:
        count = 0
        for unit in pattern:
            if unit.startswith(subject):
                count+=1
        #print(pattern)
        if count >= 1:
            students+=1
            

    proportion = students/len(data)*100
    return round(proportion,0)
    
```

```python
codes = []
for pattern in patterns:
    for unit in pattern:
        if unit[:4] not in codes:
            codes.append(unit[:4])
            

```

```python
percentage_for_each_unit = {}
for each_code in codes:
    percentage_for_each_unit[each_code]= find_proportion(each_code,pattern_strings)

```

- Repeating  Proportion


```python
students = 0
for pattern in patterns:
    count_of_unit_in_pattern = 0
    for unit in pattern:
        if unit.startswith('MATH'):
            count_of_unit_in_pattern+=1
    if count_of_unit_in_pattern >=1:
        students+=1


```

```python
def find_proportion(subject, data):

    students = 0
    for pattern in data:
        count_of_unit_in_pattern = 0
        for unit in pattern:
            if unit.startswith(subject):
                count_of_unit_in_pattern+=1
        if count_of_unit_in_pattern >=1:
            students+=1
            
    proportion = round(students/len(data),2)
    return proportion

    
```

```python
find_proportion('CITS', pattern_strings)
```

```python

```

```python
def prop(entered_code, pattern_strings):
    entered_code = entered_code.upper()

        
    codes = [] # list of valid codes 
    for pattern in pattern_strings:
        for unit in pattern:
            if unit[:4] not in codes:
                codes.append(unit[:4])
    if entered_code not in codes:
        raise ValueError('Not found, try different value')
    
    
    return find_proportion(entered_code, pattern_strings)

        
```

```python
codes = []
for pattern in patterns:
    for unit in pattern:
        if unit[:4] not in codes:
            codes.append(unit[:4])
            

            
def list_prop(enter_list, pattern):
    
    dic = {}        
    for code in enter_list:
        try:
            dic[code] = prop(code, pattern)
        except ValueError:
            pass
    
    return dic

```

```python
list_prop(["Mah", "CITS"], pattern_strings)
```

```python

```

### Challenge

* Print a list of each unit that is being studied, and how many students are studying it, in order from the unit with the most students to the unit with the least.

Hints: Use a [dictionary](https://docs.python.org/3/library/stdtypes.html#mapping-types-dict), look at the `get` method, and consider the `key` argument of the inbuilt function `sorted`.

(My code is 7 lines including the print statement.)

* How many units are being studied altogether?

```python
DATA = "unit-patterns.txt"
patterns = get_patterns(DATA)
print(patterns[:5])
```

Output:
```
[['CITS2224-1', 'CITS2023-2', 'CITS3221-2', 'CITS3224-1', 'CITS3225-1'], ['CITS2023-2', 'ENSC3221-2', 'ENSC3224-1', 'ENSC3225-1', 'ENSC3227-2', 'ENSC3210-1', 'ENSC3221-1', 'MECH3023-2'], ['CITS1021-2', 'CITS2021-1', 'CITS2023-2', 'ENSC2224-1', 'ENSC3211-2', 'ENSC3213-2', 'MATH2221-1', 'STAT1524-1'], ['CITS2023-2', 'CITS3220-2', 'CITS3223-2', 'CITS3021-1', 'CITS3025-1', 'MATH3224-1', 'MATH3221-1', 'STAT3264-1', 'STAT3260-2'], ['BIOC3221-1', 'BIOC3224-1', 'BIOC3225-2', 'BIOC3225-2', 'CITS1023-2', 'CITS2021-1', 'CITS2023-2', 'MATH1211-1']]
```

```python

```

```python
counts = {}

for i in range(len(patterns)):
    pattern = patterns[i]
    for j in range(len(pattern)):
        unit = pattern[j]
        counts[unit] = counts.get(unit, 0)+1
        
        
        



```

```python
counts = {}

for i in range(len(patterns)):
    pattern = patterns[i]
    for j in range(len(pattern)):
        unit = pattern[j]
        counts[unit] = counts.get(unit, 0)+1

```

```python
for pattern in patterns[:5]:
    print(pattern)
```

Output:
```
['CITS2224-1', 'CITS2023-2', 'CITS3221-2', 'CITS3224-1', 'CITS3225-1']
['CITS2023-2', 'ENSC3221-2', 'ENSC3224-1', 'ENSC3225-1', 'ENSC3227-2', 'ENSC3210-1', 'ENSC3221-1', 'MECH3023-2']
['CITS1021-2', 'CITS2021-1', 'CITS2023-2', 'ENSC2224-1', 'ENSC3211-2', 'ENSC3213-2', 'MATH2221-1', 'STAT1524-1']
['CITS2023-2', 'CITS3220-2', 'CITS3223-2', 'CITS3021-1', 'CITS3025-1', 'MATH3224-1', 'MATH3221-1', 'STAT3264-1', 'STAT3260-2']
['BIOC3221-1', 'BIOC3224-1', 'BIOC3225-2', 'BIOC3225-2', 'CITS1023-2', 'CITS2021-1', 'CITS2023-2', 'MATH1211-1']
```

```python
for j in range(len(patterns[:5])):
    pattern = patterns[j]
    print(pattern)
```

Output:
```
['CITS2224-1', 'CITS2023-2', 'CITS3221-2', 'CITS3224-1', 'CITS3225-1']
['CITS2023-2', 'ENSC3221-2', 'ENSC3224-1', 'ENSC3225-1', 'ENSC3227-2', 'ENSC3210-1', 'ENSC3221-1', 'MECH3023-2']
['CITS1021-2', 'CITS2021-1', 'CITS2023-2', 'ENSC2224-1', 'ENSC3211-2', 'ENSC3213-2', 'MATH2221-1', 'STAT1524-1']
['CITS2023-2', 'CITS3220-2', 'CITS3223-2', 'CITS3021-1', 'CITS3025-1', 'MATH3224-1', 'MATH3221-1', 'STAT3264-1', 'STAT3260-2']
['BIOC3221-1', 'BIOC3224-1', 'BIOC3225-2', 'BIOC3225-2', 'CITS1023-2', 'CITS2021-1', 'CITS2023-2', 'MATH1211-1']
```

```python

counts = {}
for pattern in patterns:
    for unit in pattern:
        counts[unit] = counts.get(unit,0)+1

```

# Sorting Dictionary

```python

sorted(counts.items(),key = lambda kv:kv[1],reverse = True)
```

Output:
```
[('CITS2023-2', 255),
 ('CITS1023-2', 79),
 ('CITS1021-1', 77),
 ('CITS3025-1', 74),
 ('STAT2021-1', 72),
 ('STAT2023-2', 67),
 ('STAT1024-1', 60),
 ('CITS3021-1', 53),
 ('CITS2223-2', 48),
 ('CITS3223-2', 45),
 ('CITS2224-1', 43),
 ('CITS3220-2', 43),
 ('CITS2211-2', 36),
 ('STAT3260-2', 33),
 ('CITS3221-2', 28),
 ('CITS2021-1', 27),
 ('CITS1221-1', 26),
 ('CITS3224-1', 23),
 ('STAT1023-2', 21),
 ('CITS3225-1', 17),
 ('MATH2221-1', 13),
 ('MATH1214-1', 13),
 ('ENSC2224-1', 12),
 ('FINA1129-2', 11),
 ('STAT1524-1', 10),
 ('CITS1221-2', 10),
 ('FINA2223-2', 10),
 ('ACCT1121-1', 10),
 ('ECON1124-1', 10),
 ('MATH2231-1', 10),
 ('FINA1221-1', 9),
 ('MKTG1225-1', 9),
 ('MATH1211-1', 8),
 ('MKTG1225-2', 8),
 ('ENSC2221-1', 8),
 ('MATH1721-1', 8),
 ('PHIL1224-1', 8),
 ('ECON1121-1', 8),
 ('CITS2021-2', 8),
 ('INDG1223-2AA', 8),
 ('ENSC3221-2', 7),
 ('ENSC3227-2', 7),
 ('CITS1021-2', 7),
 ('FINA3320-1', 7),
 ('PSYC1121-1', 7),
 ('MUSC1255-2', 7),
 ('STAT2263-2', 7),
 ('ENSC3211-2', 6),
 ('PHYS
```

```python

```

```python

```

- Practice Code

```python
with open(DATAFILE, 'r') as file:
    print(file.readlines)
```

```python
line = ' CITS2224-1 + CITS2023-2 + CITS3221-2 + CITS3224-1 + CITS3225-1 \t 1 (2.0%) \n'
```

```python
line.partition("\t 1 (2.0%) \n")
```

```python
line.partition('\t')
```

#### Sort Patterns
- How does "not equal" works?


```python
def sort_patterns(patterns):
    newpattern= sorted(patterns)
    return newpattern
```

```python

students = len(pattern_strings)
count = 0
sorted_file = []
while count != students:
    sorted_file.append(sorted(patterns[count]))
    count+=1
sorted_file = sorted(sorted_file)



```

```python
sorted_patterns =[]
for pattern in pattern_strings:
    sorted_patterns.append(sorted(pattern_strings))
sorted_pattern = sorted(sorted_patterns)



```

### Remove Duplicates


```python

new_patterns = [sorted_pattern[0]]
for pattern in sorted_pattern[1:]:
    if pattern != new_patterns[-1]:
        new_patterns.append(pattern)

```

```python

```

```python

```

```python

```

```python

```

### Find Unique Patterns

- sort files
- find unique patterns

```python
sorted_pattern = []
for pattern in pattern_strings:
    sorted_pattern.append(sorted(pattern))

patterns_sorted = sorted(sorted_pattern)



        
```

```python
unique_patterns = [patterns_sorted[0]]
for unique in patterns_sorted[1:]:
    if unique != patterns_sorted[-1]:
        unique_patterns.append(unique)

    
        
```

```python
uniq = []

for pattern in pattern_strings:
    if pattern not in uniq:
        uniq.append(pattern)
```

```python

```

```python
sorted_file = []
for pattern in pattern_strings:
    sorted_file.append(sorted(pattern))
sorted_file = sorted(sorted_file)



```

```python

unique_file = [sorted_file[0]]
for uniques in sorted_file[1:]:
    if uniques != unique_file[-1]:
        unique_file.append(uniques)

```

# Find Proportion


- find proportion for one subject
- Then generalise

### 1 Subject 

```python
student = 0
for pattern in unique_file:
    count = 0
    for unit in pattern:
        if unit.startswith('MATH'):
            count+=1


    if count >= 1:
        student+=1
        


proportion = round(student/len(unique_file),2)

```

```python
proportion
```

Output:
```
0.22
```

### General Units, any unit proportion

# Find the list of Unique Subjects

```python
def proportion(units,data):
    
    units = units.upper()

    student = 0
    for pattern in data:
        count = 0
        for unit in pattern:
            if unit.startswith(units):
                    count+=1
        if count >= 1:
            student+=1

    proportion = student/len(data)

    return round(proportion,2)


    
```

```python
unique_units = []
for pattern in unique_file:
    for unit in pattern:
        if unit[:4] not in unique_units:
            unique_units.append(unit[:4])

            

```

```python
def all_prop(code_list,pattern):
    dictionary = {}
    for code in code_list:
        try:
            dictionary[code] = proportion(code,pattern)
        except:
            pass
        
    return dictionary
```

```python
all_prop(["code"], unique_file)
```

Output:
```
{'code': 0.0}
```

```python

```